# Einführung Python

In dieser Einführung lernen Sie die Grundlagen der Programmiersprache Python. Doch zuerst ein paar wichtige Anmerkungen.

Ausführen einer Zelle in Jupyter Notebook:

    <Shift> + <Enter>

Ausgabe in Python:


```python
print(5)                        # Gibt den Wert 5 aus.
print("Dies ist eine Ausgabe!") # Gibt einen Satz aus.
print(5+5)                      # Gibt das Ergebnis der Rechnung aus.
print("Das Ergebnis ist:", 7*3) # Kombinierte Ausgabe aus Satz und Rechnung.
```

Kommentare in Python:

```python
# Dies ist ein Kommentar.
```

## Variablen - Datentypen und Datenstrukturen

### Datentypen

Bei dem Programmieren können Werte Variablen zugeordnet werden. Diese Variablen sind im Grunde wie Kisten. In einer Kiste kann ein Gegenstand verstaut werden. Um nach einiger Zeit noch zu wissen was in dieser Kiste verstaut worden ist, werden diese Beschriftet.
Nichts anderes wird bei den Variablen gemacht. Es wird ein Wert einer selbst benannten Variablen zugeordnet.

```python
x = 3
nummer = 40
```

    Bei der Benennung der Variablen muss darauf geachtet werden, 
    dass keine Umlaute oder Leerzeichen verwendet werden!

Je nach zugeordneten Werten, werden die Variablen in unterschiedliche Datentypen unterschieden:

- Integer (Ganzzahl)

```python
x = 100
y = -20
```

- Float (Gleitkommazahl)

```python
temperatur = 27.3
durchschnitt = 1.94
```

- String (Zeichenkette)

```python
name = "Max Mustermann"
wort = 'Buch'
```

- Bool (boolesche Werte)

```python
zustand1 = True
zustand2 = False
```

Um herauszufinden welchen Datentyp eine Variable besitzt, kann die Funktion "type(x)" verwendet werden.

```python
wort = "String"
type(wort)
```

### Datenstrukturen

Bei den bisher genannten Datentypen wurde immer genau ein Wert eine Variable zugeordnet. Es gibt jedoch Datentypen, bei denen mehrere Werte eine Variable zugeordnet werden können.

#### Liste

Einer dieser Datentypen ist die Liste. Ihr können Werte mit den oben genannten Datentypen oder auch andere Listen zugeordnet werden. Listen werden wie folgt deklariert:

```python
leer = []                   # Erstellt eine leere Liste.
liste = [1, 2, 3]
namen = ["Max", "Anna", "Simon"]
a = [liste, namen]
```

Listen können auch nach der Deklaration verändert werden. Es ist auch möglich an einzelne Werte der Liste zu gelangen. Dabei ist zu beachten, dass bei dem Programmieren von Null an gezählt wird.

```python
liste = [10, 11, 12]          # Deklarierung der Liste.
liste[0]                      # Wert der Liste an der Stelle 1 (Wert = 10)
```

Zu den Listen gehören auch diverse Methoden, die das Verarbeiten dieser vereinfacht. Im Folgenden sind einige Methoden als Beispiel aufgelistet:

```python
liste.append(x)               # Hängt den Wert x an das Ende der Liste.
liste.clear()                 # Löscht alle Elemente der Liste.
liste.copy()                  # Erstellt eine Kopie der Liste.
liste.count(x)                # Zählt wie oft der Wert x in der Liste vorhanden ist.
usw.
```

In Jupyter Notebook besteht die Möglichkeit, sich verfügbare Methoden anzeigen zu lassen. Zum Beispiel bei einer Liste: 

    liste. + <Tab>

#### Tuple

Ein Tuple ist im Gegensatz zu einer Liste unveränderlich. Ein Tuple wird wie folgt deklariert:

```python
tuple1 = 1, 2, 3
tuple2 = (10, 11, 12)           # Beide Varianten sind möglich.
```

#### Dictionaries 

Dictionaries sind wie Listen. Der Unterschied besteht darin, dass beliebige Datentypen als Indizes verwendet werden können.
Deklaration eines Dictionaries:

```python
angestellter = {"name" : "Mustermann", "vorname" : "Max", "alter" : "27", "gehalt" : 3750.63}
```

Wie oben gezeigt werden Werte den selbst gewählten Indizes zugewiesen. Auf die einzelnen Werte wird wie folgt zugegriffen:

```python
angestellter["gehalt"]       # Hat den Wert 3750.63
```

### Aufgabe 1

<!-- grade_id: cell-a8595056c533fb71, points: 0, grade: false, solution: true -->
```python
'''Gegeben ist die unten stehende Liste. Wie oft kommt in dieser der Buchstabe "f" vor?
Benutzen Sie dafür eine Listen-Funktion und speichern Sie das Ergebnis in einer Variablen mit
dem Namen anz .'''

aufgabe1 = ['a', 'y', 'h', 'f', 'i', 'x', 'e', 'f', 5, 6, 'f', 't', 45, 'f']
# YOUR CODE HERE

```

<!-- grade_id: cell-cb2ba73c6ba6331f, points: 0, grade: true, solution: false -->
```python
'''Test'''
assert anz == 4, "Ihr Lösungsweg ist Falsch"
```

## Mathematische Operatoren

In Python ist es auch möglich mit Zahlen oder Variablen zu rechnen. Neben den Grundrechenarten und der Potenzrechnung gibt es auch Module die eine Vielzahl von weiteren Methoden bereitstellen.

#### Addition

```python
x = 3
y = 2

z = x + y
s = 7 + 5
```

#### Subtraktion

```python
x = 10
y = 3

z = x - y
d = 40 - 20
```

#### Multiplikation

```python
x = 3
y = 5

z = x * y
p = 10 * 10
```

#### Division 

```python
x = 15
y = 5

z = x / y
q = 40 / 4
```

#### Potenzrechnung


```python
x = 5
y = 2

z = x ** y
e = 2 ** 2
```

#### Ganzzahlige Division

```python
x = 5
y = 2

z = x // y
e = 10 // 3 
```

#### Modulo

```python
x = 5
y = 2

z = x % y
e = 10 % 9 
```

## Module 

Wie in anderen Programmiersprachen gibt es auch in Python die Möglichkeit Bibliotheken (in Python werden diese Module genannt) hinzuzufügen. Je nach verwendeter IDE (Integradet Development Environment) sind diverse Module bereits installiert. Wenn dies nicht der Fall sein sollte, muss dies vom Benutzer manuell durchgeführt werden. Um die Methoden eines Modules nutzen zu können, muss dieses importiert werden. Dies ist wie folgt an dem Beispiel des numpy-Modules dargestellt:

```python
import numpy as np          # Importieren des Modules als np.
                                                  

np.sin(np.pi/2)
```

## Schleifen

Ein wichtiges Element bei dem Programmieren sind die Schleifen. Sie werden benötigt, wenn ein gewisser Teil des geschriebenen Codes, bis zum Erfüllen einer Bedingung, wiederholt werden soll. In Python gibt es zwei Arten von Schleifen, einmal die while-Schleife und einmal die for-Schleife. Beide Schleifen sind Kopfgesteuert (Abarbeitung der Schleife erst nach Überprüfung der Bedingung).

### Vergleichsoperatoren und Logische Operatoren

In Bedingungen werden häufig Operatoren verwendet um zwei oder mehrere Variablen zu vergleichen. Dabei wird in Vergleichsoperatoren und Logische Operatoren unterschieden.

#### Vergleichsoperatoren

|Bedeutung          |Operator|Beispiel |Ergebnis|
|   :-:             |  :-:   |  :-:    |  :-:   |
| Gleich            |   ==   | 4 == 5  |  False |
|Ungleich           |   !=   | 4 == 5  |  True  |
| Kleiner           |   <    | 10 < 20 |  True  |
| Größer            |   >    | 10 > 20 |  False |
|Kleiner oder gleich|   <=   | 15 <= 15|  True  |
|Größer oder gleich |   >=   | 23 >= 30|  False |

#### Logische Operatoren

|Bedeutung|Operator|Beispiel                |Ergebnis|
|   :-:   |  :-:   |  :-:                   |  :-:   |
|   und   |  and   | (5 == 5) and (10 == 10)|  True  |
|   oder  |  or    | (5 == 5) or (10 == 10) |  True  |
|         |        | (5 == 5) or (10 == 5)  |  True  |
|         |        | (2 == 5) or (10 == 5)  |  False |

### while-Schleife

Die while-Schleife wird häufig dann genutzt, wenn nicht bekannt ist, wie oft die Schleife durchlaufen werden soll. Sie wird wie folgt genutzt:

```python
x = 0

while(x < 5):                   # Gibt die Zahlen 1 bis 5 aus.
    x += 1
    print(x)
```

### for-Schleife

Die for-Schleife wird dann benutzt, wenn von Anfang an klar ist, wie oft die Schleife durchlaufen werden soll. 
Mit ihr können z.B. Listen durlaufen werden:

```python
namen = ['Peter', 'Klaus', 'Helmut']

for x in namen:
    print(x)
```

In Kombination mit der for-Schleife wird häufig die range-Funktion verwendet. Diese Funktion gibt eine Reihe von Zahlen aus, deren Anfangswert, Endwert und Schrittweite durch den Nutzer festgelegt wird.

```python
x = 0      # Anfangswert
y = 6      # Endwert +1 !
z = 1      # Schrittweite

for i in range(x, y, z):
    print(i)
```

Die Übergabe des Startwertes und Schrittweite ist optional. Werden keine Werte übergeben sind diese 0 für den Startwert und 1 für die Schrittweite.

```python
for i in range(4):
    print(i)
```

### Aufgabe 2

<!-- grade_id: cell-f4e5286e71603910, points: 0, grade: false, solution: true -->
```python
'''Schreiben Sie eine Schleife, welche die Zahlen von 2 bis 10, in zweier Schritten, in die
 unten stehende Liste schreibt.'''

aufgabe2 = []
# YOUR CODE HERE
```

<!-- grade_id: cell-520006f2613ce2f2, points: 0, grade: true, solution: false -->
```python

```

## Funktionen

Beim Programmieren kommt es vor, dass gewisse Codesequenzen mehr als einmal ausgeführt werden müssen. Um nicht jedes Mal den besagten Code erneut zu schreiben, kann dieser einmalig als Funktion definiert werden. Dies hat den Vorteil, dass das Programm übersichtlicher wird.

Eine Funktion besteht aus einem Funktionsnamen, Übergabeparameter und dem Code, der ausgeführt werden soll. Es können Funktionen so gestaltet werden, dass sie einen Rückgabewert haben oder nicht. Sie werden wie folgt definiert:

```python
def verdoppeln(parameter):          # Funktion mit Rückgabewert
    
    x = parameter * 2
    return x

def ausgabe(name):                  # Funktion ohne Rückgabewert
    
    print("Vorname: " + name)
    

produkt = verdoppeln(2)             # Verwendung der Funktionen
print(produkt)
ausgabe("Max")
```

In der Regel werden die Namen der Funktionen so gewählt, dass direkt ersichtlich ist, was diese tut.

### Aufgabe 3

<!-- grade_id: cell-af6d295a19d02de6, points: 0, grade: false, solution: true -->
```python
'''Schreiben Sie eine Funktion, welche den Flächeninhalt eines Dreiecks zurückgibt.
    Die Funktion soll wie folgt benannt werden: flaecheDreieck(g, h)'''
# YOUR CODE HERE
```

<!-- grade_id: cell-9f035a3fc65b3c20, points: 0, grade: true, solution: false -->
```python
'''Test'''
assert flaecheDreieck(10, 10) == 50, "Ihr Lösungsweg ist Falsch!"
assert flaecheDreieck(20, 20) == 200, "Ihr Lösungsweg ist Falsch!"
```

## Klassen


Bei objektorientierten Programmiersprachen wie Python gibt es die Möglichkeit Klassen zu erstellen. Dabei sind Objekte, bestehend aus Attributen und Funktionen, in einer Klasse beschrieben. Die Funktionen der Klasse werden Methoden genannt. Klassen vereinfachen das Erweitern des Programmes, da es schnell möglich ist mehrere Objekte zu erzeugen.  Der Name einer Klasse ist frei wählbar, wobei in der Regel der erste Buchstabe großgeschrieben wird. Dies erleichtert das Erkennen von Klassen im Programm.

Folgend ist Anhand eines Beispiels das Definieren einer Klasse gezeigt. Dabei ist die erste Methode ein Konstruktor. Dieser ist für das Zuweisen von Werten zu den im ihn Definierten Attributen zuständig. Der Konstruktor wird direkt nach dem Erstellen eines neuen Objektes aufgerufen. Die zweite Methode ist ein einfaches Ausgeben der Attribute eines Objektes.


```python
class Person: # Definiert werden Klassen mit dem Schlüsselwort class gefolgt von dem Klassennamen.
    
    def __init__(self, vornameZuw = "-", nachnameZuw = "-", alterZuw = "-"): # Dies ist der Konstruktor. 
                                                                             # In ihm werden auch die Standardwerte definiert.
        self.vorname = vornameZuw          # Das self. gibt den Konstruktor an, dass die Attribute Klassenweit definiert sind.
        self.nachname = nachnameZuw
        self.alter = alterZuw
        
    def ausgabe(self): 
        print(self.vorname)
        print(self.nachname)
        print(self.alter)
        print("-------------------")
```

Die Objekte werden wie folgt erstellt, benutzt und bearbeitet:

```python
p1 = Person()                              # Objekt mit den Standardwerten.
p2 = Person("Max", "Mustermann", 25)       # Objekt mit zugewiesenen Werten.
p3 = Person("Anna", "Musterfrau", 27)


p1.ausgabe()                               # Nutzen einer Methode.
p2.ausgabe()
p3.ausgabe()

p1.vorname = "Peter"                       # Es können auch später die Werte geändert werden.
p1.nachname = "Müller"
p1.alter = 30
```